# Basic Python assignment

## Write code to sum all elements of a list

<!-- grade_id: cell-336655a6e374838c, points: 0, grade: false, solution: true -->
```python
def sumList(lst):
    # YOUR CODE HERE
    return sum(lst)
```

<!-- grade_id: cell-38e3a2715d4a9712, points: 5, grade: true, solution: false -->
```python
lst = []
with open("/mnt/nfs/dcdemo110@gmail.com/share/testcase.txt") as file:
    for line in file:
        line = line.strip() #preprocess line
        lst.append(int(line))
assert(sumList(lst[:-1]) == lst[-1])
```

# Plot a graph using matplotlib showing how a normal(0,1) looks like

<!-- grade_id: cell-d662ec1f00f9bb61, points: 10, grade: true, solution: true -->
```python
import matplotlib.pyplot as plt
import numpy as np
import matplotlib.mlab as mlab
import math
%matplotlib inline
mu = 0
variance = 1
sigma = 1

x = np.linspace(mu - sigma, mu + sigma, 100)
plt.plot(x,mlab.normpdf(x, mu, sigma))
plt.show()
```